# Mandatory Assignment 1

This is the first of two mandatory assignments which must be completed during the course. First some practical information:

* When is the assignment due?: **23:59, Sunday, August 19, 2018.**
* How do you grade the assignment?: You will **peergrade** each other as primary grading. 
* Can i work with my group?: **yes**

The assigment consist of one to tree problems from each of the exercise sets you have solved so far (excluding Set 1). We've tried to select problems which are self contained, but it might be nessecary to solve some of the previous exercises in each set to fully answer the problems in this assignment.

## Problems from Exercise Set 2:

> **Ex. 2.2**: Make two lists. The first should be numbered. The second should be unnumbered and contain at least one sublevel. 

# [Answer to Ex. 2.2 here] (convert to markdown cell)

### First list

1. Item number 1
1. Item number 2
1. Item number 3

### Second list

* Item number 1
* Item number 2
    * Item number 2a
    * Item number 2b


## Problems from Exercise set 3:

> **Ex. 3.1.3:** Let `l1 = ['r ', 'Is', '>', ' < ', 'g ', '?']`. Create from `l1` the sentence `"Is r > g?"` using your knowledge about string formatting. Store this new string in a variable called `answer_31`. Make sure there is only one space in between worlds.
>
>> _Hint:_ You should be able to combine the above informations to solve this exercise.

<!-- grade_id: problem_31, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 3.1.3 here]
l1 = ['r ', 'Is', '>', ' < ', 'g ', '?']

# answer_31 = 

# YOUR CODE HERE

answer_31 = l1[1].strip() + ' ' + l1[0].strip() + ' ' + l1[2].strip() + ' ' + l1[4].strip() + l1[5].strip()

```

<!-- grade_id: problem_31_tests, points: 1, grade: true, solution: false -->
```python
assert answer_31 == "Is r > g?"
```

> **Ex. 3.1.4**: Create an empty dictionary `words` using the `dict()`function. Then add each of the words in `['animal', 'coffee', 'python', 'unit', 'knowledge', 'tread', 'arise']` as a key, with the value being a boolean indicator for whether the word begins with a vowel. The results should look like `{'bacon': False, 'asynchronous': True ...}`. Store the result in a new variable called `answer_32`.
>
>> _Hint:_ You might want co first construct a function that asseses whether a given word begins with a vowel or not.

<!-- grade_id: problem_32, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 3.1.4 here]
W = ['animal', 'coffee', 'python', 'unit', 'knowledge', 'tread', 'arise']

# answer_32 = 


# YOUR CODE HERE

def VowCheck (z):
    if(z=='A' or z=='a' or z=='E' or z =='e' or z=='I' or z=='i' or z=='O' or z=='o' or z=='U' or z=='u'):
        return True
    else:
        return False
    
answer_32 = dict()

keys = ['animal', 'coffee', 'python', 'unit', 'knowledge', 'tread', 'arise']
values = [VowCheck(keys[0][0]), VowCheck(keys[1][0]), VowCheck(keys[2][0]), 
          VowCheck(keys[3][0]), VowCheck(keys[4][0]), VowCheck(keys[5][0]), VowCheck(keys[6][0])]
key_value_pairs = list(zip(keys, values, ))

answer_32.update (key_value_pairs)

print(answer_32)

```

Output:
```
{'animal': True, 'coffee': False, 'python': False, 'unit': True, 'knowledge': False, 'tread': False, 'arise': True}
```

<!-- grade_id: problem_32_tests, points: 1, grade: true, solution: false -->
```python
assert answer_32 == {i: i[0] in 'aeiou' for i in W}
assert sorted(answer_32) == sorted(W)
```

> **Ex. 3.3.2:** use the `requests` module (get it with `pip install requests`) and `construct_link()` which you defined in the previous question (ex 3.3.1) to request birth data from the "FOD" table. Get all available years (variable "Tid"), but only female births (BARNKON=P) . Unpack the json payload and store the result. Wrap the whole thing in a function which takes an url as input and returns the corresponding output.
>
> Store the birth data in a new variable called `answer_33`.
>
>> _Hint:_ The `requests.response` object has a `.json()` method. 
>
>> _Note:_ you wrote `construct_link()` in 3.3.1, if you didn't heres the link you need to get: `https://api.statbank.dk/v1/data/FOLK1A/JSONSTAT?lang=en&Tid=*`

<!-- grade_id: problem_33, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 3.3.2 here]

# answer_33 = 

# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-6-eb6cf3528aa9> in <module>
      4 
      5 # YOUR CODE HERE
----> 6 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: problem_33_tests, points: 1, grade: true, solution: false -->
```python
assert sorted(answer_33['dataset'].keys()) == ['dimension', 'label', 'source', 'updated', 'value']
assert 'BARNKON' in answer_33['dataset']['dimension'].keys()
```

## Problems from exercise set 4


```python
import numpy as np
import pandas as pd
```

> **Ex. 4.1.1:** Use Pandas' CSV reader to fetch  daily data weather from 1864 for various stations - available [here](https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/). Store the dataframe in a variable called `answer_41`.
>
>> *Hint 1*: for compressed files you may need to specify the keyword `compression`.
>
>> *Hint 2*: keyword `header` can be specified as the CSV has no column names.
>
>> *Hint 3*: Specify the path, as the URL linking directly to the 1864 file. 

<!-- grade_id: problem_41, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.1 here]

# YOUR CODE HERE

path_1864 = 'https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/1864.csv.gz'

answer_41 = pd.read_csv(path_1864, compression = 'gzip', header = None)


```

<!-- grade_id: problem_41_tests, points: 1, grade: true, solution: false -->
```python
assert answer_41.shape == (27349, 8)
assert list(answer_41.columns) == list(range(8))
```

> **Ex. 4.1.2:** Structure your weather DataFrame by using only the relevant columns (station identifier, data, observation type, observation value), rename them. Make sure observations are correctly formated (how many decimals should we add? one?).
>
> Store the resulting dataframe in a new variable called `answer_42`.
>
>> *Hint:* rename can be done with `df.columns=COLS` where `COLS` is a list of column names.

<!-- grade_id: problem_42, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.2 here]

# YOUR CODE HERE

answer_42 = answer_41.drop(columns = [4, 5,6, 7], axis = 1)

answer_42 = answer_42.rename(columns = {0: 'Station Identifier', 1: 'Date', 2: 'Observation type', 3: 'Observation value'})

answer_42['Observation value'] = answer_42['Observation value']/10
#One decimal is efficient

print(answer_42.head())

```

<!-- grade_id: problem_42_tests, points: 1, grade: true, solution: false -->
```python
assert answer_42.shape == (27349, 4)
assert 144.8 in [answer_42[i].max() for i in answer_42]
assert -666.0 in [answer_42[i].min() for i in answer_42]
assert 18640101 in [answer_42[i].min() for i in answer_42]
```

> **Ex. 4.1.3:**  Select data for the station `ITE00100550` and only observations for maximal temperature. Make a copy of the DataFrame. Explain in a one or two sentences how copying works.
>
> Store the subsetted dataframe in a new variable called `answer_43`.
>
>> *Hint 1*: the `&` operator works elementwise on boolean series (like `and` in core python).
>
>> *Hint 2*: copying of the dataframe is done with the `copy` method for DataFrames.

<!-- grade_id: problem_43, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.3 here]


# YOUR CODE HERE

answer_43 = answer_42[answer_42['Station Identifier'].str.contains("ITE00100550") & answer_42['Observation type'].str.contains("TMAX")]

print(answer_43.head(10))

```

<!-- grade_id: problem_43_tests, points: 1, grade: true, solution: false -->
```python
assert 'ITE00100550' in [answer_43[i].min() for i in answer_43]
assert 'ITE00100550' in [answer_43[i].max() for i in answer_43]
assert 'TMAX' in [answer_43[i].min() for i in answer_43]
assert 'TMAX' in [answer_43[i].max() for i in answer_43]
```

> **Ex. 4.1.4:** Make a new column in `answer_44` called `TMAX_F` where you have converted the temperature variables to Fahrenheit. Make sure not to overwrite `answer_43`.
>
> Store the resulting dataframe in a variable called `answer_44`.
>
>> *Hint*: Conversion is $F = 32 + 1.8*C$ where $F$ is Fahrenheit and $C$ is Celsius.

<!-- grade_id: problem_44, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.4 here]
answer_44 = answer_43.copy()

# YOUR CODE HERE

answer_44['TMAX_F'] = 32 + 1.8 * answer_44['Observation value']

print(answer_44.head(10))

```

<!-- grade_id: problem_44_tests, points: 1, grade: true, solution: false -->
```python
assert set(answer_44.columns) -  set(answer_43.columns) == {'TMAX_F'}
```

## Problems from exercise set 5

```python
import matplotlib.pyplot as plt
import numpy as np 
import pandas as pd
import seaborn as sns 

%matplotlib inline 

iris = sns.load_dataset('iris')
titanic = sns.load_dataset('titanic')
```

> **Ex. 5.1.1:**: Show the first five rows of the titanic dataset. What information is in the dataset? Use a barplot to show the probability of survival for men and women within each passenger class. Can you make a boxplot showing the same information (why/why not?). _Bonus:_ show a boxplot for the fare-prices within each passenger class. 
>
> Spend five minutes discussing what you can learn about the survival-selection aboard titanic from the figure(s).
>
> > _Hint:_ https://seaborn.pydata.org/generated/seaborn.barplot.html, specifically the `hue` option.


<!-- grade_id: problem_51, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.1 here]

# YOUR CODE HERE

print(titanic.head(5))

sns.barplot(x = titanic['class'], y = titanic['survived'], hue = titanic['sex'])

```

```python
t2 = sns.boxplot(x='class', y='fare', hue='sex', data=titanic)
```

> **Ex. 5.1.2:** Using the iris flower dataset, draw a scatterplot of sepal length and petal length. Include a second order polynomial fitted to the data. Add a title to the plot and rename the axis labels.
> _Discuss:_ Is this a meaningful way to display the data? What could we do differently?
>
> For a better understanding of the dataset this image might be useful:
> <img src="iris_pic.png" alt="Drawing" style="width: 200px;"/>
>
>> _Hint:_ use the `.regplot` method from seaborn. 

<!-- grade_id: cell-8631e73e3796008b, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.2 here]

# YOUR CODE HERE

print(iris.head(5))

iris_plot = sns.regplot('petal_length', 'sepal_length', data = iris, scatter = True, fit_reg = True, order = 2)

iris_plot.set(xlabel = 'petal', ylabel = 'sepal')
iris_plot.set_title('Iris')

```

> **Ex. 5.1.3:** Combine the two of the figures you created above into a two-panel figure similar to the one shown here:
> <img src="Example.png" alt="Drawing" style="width: 600px;"/>
>
> Save the figure as a png file on your computer. 
>> _Hint:_ See [this question](https://stackoverflow.com/questions/41384040/subplot-for-seaborn-boxplot) on stackoverflow for inspiration.

<!-- grade_id: cell-686e1fcbc475435a, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.3 here]

# YOUR CODE HERE

f, axes = plt.subplots(1, 2, figsize = (12, 5))

sns.regplot('petal_length', 'sepal_length', data = iris, scatter = True, fit_reg = True, order = 2, ax = axes[0])
ax1 = sns.barplot(x = titanic['class'], y = titanic['survived'], hue = titanic['sex'], ax = axes[1])

ax1.set_ylim([0,1])

```

> **Ex. 5.1.4:** Use [pairplot with hue](https://seaborn.pydata.org/generated/seaborn.pairplot.html) to create a figure that clearly shows how the different species vary across measurements. Change the color palette and remove the shading from the density plots. _Bonus:_ Try to explain how the `diag_kws` argument works (_hint:_ [read here](https://stackoverflow.com/questions/1769403/understanding-kwargs-in-python))

<!-- grade_id: cell-0f5793c056a02040, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.4 here]

# YOUR CODE HERE

ax_514 = sns.pairplot(iris, hue = "species", palette = "husl",  diag_kws=dict(shade=False))
ax_514

```

## Problems from exercise set 6

> _Note:_ In the exercises we asked you to download weather data from the NOAA website. For this assignment the data are loaded in the following code cell into two pandas dataframes.

```python
%matplotlib inline
import pandas as pd
import matplotlib.pyplot as plt

weather_1864 = pd.read_csv('weather_data_1864.csv')
```

Output:
```
pandas/_libs/parsers.pyx in pandas._libs.parsers.TextReader.__cinit__()

pandas/_libs/parsers.pyx in pandas._libs.parsers.TextReader._setup_parser_source()

FileNotFoundError: [Errno 2] File b'weather_data_1864.csv' does not exist: b'weather_data_1864.csv'
```

> **Ex. 6.1.4:** Extract the country code from the station name into a separate column.
>
> Create a new column in `weather_1864` called `answer_61` and store the country codes here.
>
>> _Hint:_ The station column contains a GHCND ID, given to each weather station by NOAA. The format of these ID's is a 2-3 letter country code, followed by a integer identifying the specific station. A simple approach is to assume a fixed length of the country ID. A more complex way would be to use the [`re`](https://docs.python.org/2/library/re.html) module.

<!-- grade_id: problem_61, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 6.1.4]
# weather_1864['answer_61'] =

# YOUR CODE HERE

path_1864 = 'https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/1864.csv.gz'

df_csv = pd.read_csv(path_1864, compression = 'gzip', header = None)

df_csv = df_csv.drop(columns = [4, 5,6, 7], axis = 1)

df_csv = df_csv.rename(columns = {0: 'Station Identifier', 1: 'Date', 2: 'Observation type', 3: 'Observation value'})

df_csv['answer_61'] = df_csv['Station Identifier'].str[0:2]

df_csv['Datetime'] = df_csv['Date']

df_csv['Datetime'] = df_csv['Datetime'].astype(str)

df_csv['Datetime'] = pd.to_datetime(df_csv['Datetime'])

df_csv['month'] = pd.DatetimeIndex(df_csv['Datetime']).month

df_csv_TMAX = df_csv[df_csv['Observation type'].str.contains("TMAX")]

weather_1864 = df_csv_TMAX.copy()

```

```python
#print(weather_1864.head())

sorted(weather_1864['answer_61'].str[:2].unique())
```

<!-- grade_id: problem_61_tests, points: 1, grade: true, solution: false -->
```python
assert sorted(weather_1864['answer_61'].str[:2].unique()) == sorted(['SZ', 'CA', 'EZ', 'GM', 'AU', 'IT', 'BE', 'UK', 'EI', 'AG', 'AS'])
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-7-5b00cb64306a> in <module>
----> 1 assert sorted(weather_1864['answer_61'].str[:2].unique()) == sorted(['SZ', 'CA', 'EZ', 'GM', 'AU', 'IT', 'BE', 'UK', 'EI', 'AG', 'AS'])

NameError: name 'weather_1864' is not defined
```

> **Ex. 6.1.5:** Make a function that downloads and formats the weather data according to previous exercises in Exercise Section 4.1, 6.1. You should use data for ALL stations but still only select maximal temperature. _Bonus:_ To validate that your function works plot the temperature curve for each country in the same window. Use `plt.legend()` to add a legend. 
>
> Name your function `prepareWeatherData`.

<!-- grade_id: problem_62, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 6.1.5]

#def prepareWeatherData(year):
    # Your code here
  #  return 

# YOUR CODE HERE

def prepareWeatherData(year):
    df = pd.read_csv('https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/'+ year +'.csv.gz', compression='gzip', header=None, sep=',', quotechar='"')
    df_dropped = df.drop(columns=[4, 5, 6, 7])
    df_renamed = df_dropped.rename(columns={0: "Station identifier", 1: "Date", 2: "Observation type", 3: "Observation value"})
    df_reduced = df_renamed[df_renamed['Observation type'].str.contains("TMAX")]
    df_reduced['Fahrenheit'] = df_reduced['Observation value'] * 1.8 + 32
    df_sorted = df_reduced.sort_values(by=['Observation value'])
    df_reindexed = df_sorted.reset_index(drop=True)
    df_string = df_reindexed.astype(str)
    df_string['date_time'] = pd.to_datetime(df_string['Date'])
    df_string['Observation value'] = pd.to_numeric(df_string['Observation value'])/10
    df_string['dt_Month'] = pd.DatetimeIndex(df_string['date_time']).month
    df_string['Country code'] = df_string['Station identifier'].str.replace('\d+', '')
    df_string = df_string.drop(columns=['Date'])
    df_string = df_string.set_index('date_time')
    return(df_string)

prepareWeatherData('1864')

```

```python
f, axes = plt.subplots(1, figsize = (16, 9))

prepareWeatherData('1864').groupby('Country code')['Observation value'].plot(legend=True)

plt.legend(loc='upper center', bbox_to_anchor=(0.5, -0.1), fancybox=True, shadow=True, ncol=5)
```

<!-- grade_id: problem_62_tests, points: 1, grade: true, solution: false -->
```python
assert prepareWeatherData('1864').shape == (5686, 6)
```

## Problems from exercise set 7

> _Note:_ Once again if you haven't managed to download the data from NOAA, you can refer to the github repo to get csv-files containing the required data.

```python
%matplotlib inline 

import pandas as pd 
import matplotlib.pyplot as plt
import matplotlib as mpl

# Increases the plot size a little
mpl.rcParams['figure.figsize'] = 11, 6
```

> **Ex. 7.1.1:** Plot the monthly max,min, mean, first and third quartiles for maximum temperature for our station with the ID _'ITE00100550'_ in 1864. 

> *Hint*: the method `describe` computes all these measures.

<!-- grade_id: problem_71, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 7.1.1]

# YOUR CODE HERE

df711 = prepareWeatherData('1864')[prepareWeatherData('1864')['Station identifier'].str.contains("ITE00100550")]

df_g = (df711['Observation value'].groupby(df711['dt_Month']).describe())
df_g = df_g.loc [ :, ['max', 'mean', 'min', '25%', '75%']]

print(df_g.head(12))

```

```python
plt.boxplot(df_g)
```

> **Ex. 7.1.2:** Get the processed data from years 1864-1867 as a list of DataFrames. Convert the list into a single DataFrame by concatenating vertically. 
>
> Name the concatenated data `answer_72`

<!-- grade_id: problem_72, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 7.1.2]

# YOUR CODE HERE
    
answer_72 = []

for i in range(1864, 1868):
    path_ = 'https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/'+ str(i) +'.csv.gz'
    df_csvloop = pd.read_csv(path_, compression = 'gzip', header = None)
    df_csvloop = df_csvloop.drop(columns = [4, 5,6, 7], axis = 1)
    df_csvloop = df_csvloop.rename(columns = {0: 'Station Identifier', 1: 'Date', 2: 'Observation type', 3: 'Observation value'})
    df_csvloop['Country Code'] = df_csvloop['Station Identifier'].str.extract(r'([A-z]+)') #Use of Regex
    df_csvloop['Datetime'] = df_csvloop['Date']
    df_csvloop['Datetime'] = df_csvloop['Datetime'].astype(str)
    df_csvloop['Datetime'] = pd.to_datetime(df_csvloop['Datetime'])
    df_csvloop['month'] = pd.DatetimeIndex(df_csvloop['Datetime']).month
    df_csvloop = df_csvloop[df_csvloop['Observation type'].str.contains("TMAX")]
    answer_72.append(df_csvloop)


```

```python
answer_72 = pd.concat(answer_72, axis = 0, sort = False)

print(answer_72)
```

Output:
```
Station Identifier      Date Observation type  Observation value  \
0            ITE00100550  18640101             TMAX                 10   
8            ASN00086071  18640101             TMAX                214   
15           CA006158350  18640101             TMAX                 11   
24           EIE00101859  18640101             TMAX                 94   
28           BE000006447  18640101             TMAX                -23   
34           GME00125218  18640101             TMAX                -81   
39           EZE00100082  18640101             TMAX                -92   
42           ITE00100554  18640101             TMAX                 17   
47           ASN00066062  18640101             TMAX                244   
53           AU000005901  18640101             TMAX                -50   
55           GM000004204  18640101             TMAX                -68   
59           UK000047811  18640101             TMAX                 59   
67           UK000056225  18640101          
```

<!-- grade_id: problem_72_tests, points: 1, grade: true, solution: false -->
```python
assert answer_72.shape == (30003, 7)
```

> **Ex. 7.1.3:** Parse the station location data which you can find at https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/ghcnd-stations.txt. Merge station locations onto the weather data spanning 1864-1867.  
>
> Store the merged data in a new variable called `answer_73`.
>
> _Hint:_ The location data have the folllowing format, 

```
------------------------------
Variable   Columns   Type
------------------------------
ID            1-11   Character
LATITUDE     13-20   Real
LONGITUDE    22-30   Real
ELEVATION    32-37   Real
STATE        39-40   Character
NAME         42-71   Character
GSN FLAG     73-75   Character
HCN/CRN FLAG 77-79   Character
WMO ID       81-85   Character
------------------------------
```

> *Hint*: The station information has fixed width format - does there exist a pandas reader for that?

```python
statlocpath = 'https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/ghcnd-stations.txt'

path_2 = 'https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/1864.csv.gz'
df_csvloop2 = pd.read_csv(path_2, compression = 'gzip', header = None)
df_csvloop2 = df_csvloop2.drop(columns = [4, 5,6, 7], axis = 1)
df_csvloop2 = df_csvloop2.rename(columns = {0: 'Station Identifier', 1: 'Date', 2: 'Observation type', 3: 'Observation value'})
df_csvloop2['Country Code'] = df_csvloop2['Station Identifier'].str.extract(r'([A-z]+)') #Use of Regex
df_csvloop2['Datetime'] = df_csvloop2['Date']
df_csvloop2['Datetime'] = df_csvloop2['Datetime'].astype(str)
df_csvloop2['Datetime'] = pd.to_datetime(df_csvloop2['Datetime'])
df_csvloop2['month'] = pd.DatetimeIndex(df_csvloop2['Datetime']).month
df_csvloop2 = df_csvloop2[df_csvloop2['Observation type'].str.contains("TMAX")]
```

<!-- grade_id: problem_73, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 7.1.3]

# YOUR CODE HERE

columns = [(1,11), (13,20), (22,30), (32,37), (38,40), (42,71), (73,75), (77,79), (81,82)]

df_merged = pd.read_fwf(statlocpath, header = None, colspecs = columns)
```

```python
print(df_csvloop2)
```

Output:
```
Station Identifier      Date Observation type  Observation value  \
0            ITE00100550  18640101             TMAX                 10   
8            ASN00086071  18640101             TMAX                214   
15           CA006158350  18640101             TMAX                 11   
24           EIE00101859  18640101             TMAX                 94   
28           BE000006447  18640101             TMAX                -23   
34           GME00125218  18640101             TMAX                -81   
39           EZE00100082  18640101             TMAX                -92   
42           ITE00100554  18640101             TMAX                 17   
47           ASN00066062  18640101             TMAX                244   
53           AU000005901  18640101             TMAX                -50   
55           GM000004204  18640101             TMAX                -68   
59           UK000047811  18640101             TMAX                 59   
67           UK000056225  18640101          
```

```python
df_merged = df_merged.rename(columns = {0: 'Station Identifier', 1: 'Latitude', 2: 'Longitude', 3: 'Elevation', 4: 'State',\
                                        5: 'Name', 6: 'GSN Flag', 7: 'HCN/CRN Flag', 8: 'WMO ID'})
answer_73 = pd.merge(df_csvloop2, df_merged, how = 'left', on = 'Station Identifier')

print(answer_73)
```

Output:
```
Station Identifier      Date Observation type  Observation value  \
0           ITE00100550  18640101             TMAX                 10   
1           ASN00086071  18640101             TMAX                214   
2           CA006158350  18640101             TMAX                 11   
3           EIE00101859  18640101             TMAX                 94   
4           BE000006447  18640101             TMAX                -23   
5           GME00125218  18640101             TMAX                -81   
6           EZE00100082  18640101             TMAX                -92   
7           ITE00100554  18640101             TMAX                 17   
8           ASN00066062  18640101             TMAX                244   
9           AU000005901  18640101             TMAX                -50   
10          GM000004204  18640101             TMAX                -68   
11          UK000047811  18640101             TMAX                 59   
12          UK000056225  18640101             TMAX      
```

<!-- grade_id: problem_73_tests, points: 1, grade: true, solution: false -->
```python
assert answer_73.shape == (5686, 15)
```

## Problems from exercise set 8

> **Ex. 8.1.2.:** Use the `request` module to collect the first page of job postings.
>
> Store the response.json() object in a new variable called `answer_81`.
> 

<!-- grade_id: problem_81, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 8.1.2]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_81_tests, points: 1, grade: true, solution: false -->
```python
assert sorted(answer_81.keys()) == sorted(['Expression', 'Facets', 'JobPositionPostings', 'TotalResultCount'])
```

> **Ex. 8.1.3.:** Store the 'TotalResultCount' value for later use. Also create a dataframe from the 'JobPositionPostings' field in the json. Name this dataframe `answer_82`.

<!-- grade_id: problem_82, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 8.1.3]
# answer_82 = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_82_tests, points: 1, grade: true, solution: false -->
```python
assert answer_82.shape == (20,44)
```